<!DOCTYPE html>
<html  lang="en" >
<head>
<meta charset="utf-8"/>
<title>CSS3 Selectors: [lang="az-Arab-IR"], lang="az-IR"</title>
<link rel='author' title='Richard Ishida' href='mailto:ishida@w3.org'>
<link rel='help' href='http://www.w3.org/TR/selectors/#attribute-selectors'>
<script src="../../../resources/testharness.js"></script>
<script src="../../../resources/testharnessreport.js"></script>
<link href="../../../resources/testharness.css" rel="stylesheet" type="text/css">
<meta name="assert" content="A lang= value with language, script and region subtags will NOT match a lang attribute value with the script subtag missing.">
<style type='text/css'>
.test div { width: 50px; }
#box[lang='az-Arab-IR'] { width: 100px; }
</style>
</head>
<body>



<div class="test"><div id="box" lang="az-IR">&#xA0;</div></div>


<script>
test(function() {
assert_equals(document.getElementById('box').offsetWidth, 50);
}, " ");
</script>

<div id='log'></div>

</body>
</html>
